### Codio Activity 18.6: Naive Bayes Algorithm

**Expected Time = 60 minutes** 

**Total Points = 35** 

This activity focuses on the implementation of the Naive Bayes algorithm.  You will use the scikit-learn estimator together with your earlier vectorization strategies to model the WhatsApp text and compare to your earlier work with Logistic Regression.   

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)
- [Problem 4](#-Problem-4)
- [Problem 5](#-Problem-5)

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.naive_bayes import MultinomialNB
from sklearn.pipeline import Pipeline
from sklearn.feature_extraction.text import CountVectorizer, TfidfVectorizer
from sklearn.model_selection import train_test_split, GridSearchCV
```

[Back to top](#-Index)

### Problem 1

#### Small Example

**10 Points**

The example below is adapted from Marsland's *Machine Learning an Algorithmic Perspective*.  A small dataset where the features are whether or not a student has a looming deadline, if there is a party going on, and whether or not the student feels lazy.  The activity column is the target, and your aim is to use the naive bayes formula below:

$$P(C_i) \prod_{k} P(X_j^k = a_k | C_i)$$

```python
deadline = ['urgent','urgent','near', 'none', 'none', 'none', 'near', 'near', 'near','urgent']
party = ['yes', 'no', 'yes', 'yes', 'no', 'yes', 'no', 'no', 'yes', 'no']
lazy = ['yes', 'yes', 'yes', 'no', 'yes', 'no', 'no', 'yes', 'yes', 'no']
activity = ['party', 'study', 'party', 'party', 'pub', 'party', 'study', 'tv', 'party', 'study']
```

```python
df = pd.DataFrame({'deadline': deadline, 
                  'party': party,
                  'lazy': lazy,
                  'activity': activity})
df
```

Output:
```
deadline party lazy activity
0   urgent   yes  yes    party
1   urgent    no  yes    study
2     near   yes  yes    party
3     none   yes   no    party
4     none    no  yes      pub
5     none   yes   no    party
6     near    no   no    study
7     near    no  yes       tv
8     near   yes  yes    party
9   urgent    no   no    study
```

Here, $C_i$ represents the class in the `activity` columm.  Accordingly, if we want to predict a category of activity given the input: 

```
deadline = near
party = no
lazy = yes
```

This means we need four probabilities:

- $P(party) \times P(near | party) \times P(no party | party) \times P(lazy | party)$
- $P(study) \times P(near | study) \times p(noparty | study) \times P(lazy | study)$
- $P(pub) \times P(near | pub) \times P(noparty | pub) \times P(lazy | pub)$
- $P(tv) \times P(near | tv) \times P(noparty | tv) \times P(lazy |tv)$

Compute these four probabilities and assign them to the list `probs` in the order above (party, study, pub, tv). 

Hint: No need to calculate the probabilities by hand.

<!-- grade_id: cell-9bc8a1669a0b2d75, points: 0, grade: false, solution: true -->
```python
### GRADED
probs = []

    
### BEGIN SOLUTION
probs = [1/2*2/5*0, 3/10*1/3*1*1/3, 1/10*0, 1/10*1*1*1]
### END SOLUTION

### ANSWER CHECK
print(probs)
```

Output:
```
[0.0, 0.03333333333333333, 0.0, 0.1]
```

<!-- grade_id: cell-1654be87992e8f0d, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
probs_ = [1/2*2/5*0, 3/10*1/3*1*1/3, 1/10*0, 1/10*1*1*1]

#
#
#
#
assert probs == probs_
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### MAP solution

**10 Points**

Using these probabilities, the maximum aposteriori solution involves selecting the outcome that is associated with the highest probability.  Use your list of probabilities to identify the `argmax`.  Note you can use `np.argmax` for this or just inspect the values.  What is the activity associated with the MAP solution?  Assign your answer as a string -- `party`, `study`, `pub`, or `tv` -- to `ans2` below.

<!-- grade_id: cell-256e274715d26b4a, points: 0, grade: false, solution: true -->
```python
### GRADED
ans2 = ''

    
### BEGIN SOLUTION
ans2 = 'tv'
### END SOLUTION

### ANSWER CHECK
print(ans2)
```

Output:
```
tv
```

<!-- grade_id: cell-d973b12508aa9d50, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans2_ = 'tv'

#
#
#
#
assert ans2 == ans2_
### END HIDDEN TESTS
```

### Larger Example

Now, you are to use the scikitlearn vectorizers together with the `MultinomialNB` estimator to implement naive bayes algorithm for classifying the WhatsApp data.  The data is loaded and split for you below.

```python
happy_df = pd.read_csv('data/Emotion(happy).csv')
sad_df = pd.read_csv('data/Emotion(sad).csv.zip', compression = 'zip')
full_df = pd.concat([happy_df, sad_df]).reset_index(drop = True)
X = full_df.drop('sentiment', axis = 1)
y = full_df['sentiment']
X_train, X_test, y_train, y_test = train_test_split(X['content'], y, random_state = 42)
```

[Back to top](#-Index)

### Problem 3

#### Pipeline with `CountVectorizer`

**5 Points**

Below, create a pipeline called `cvect_pipe` with named steps `cvect` and `bayes` that first vectorizes the text and then uses the `MultinomialNB` estimator with all default settings.  Fit this on the train and score it on the test, assigning the accuracy to `cvect_acc` below.

<!-- grade_id: cell-6a6f73de4a0f6334, points: 0, grade: false, solution: true -->
```python
### GRADED
cvect_pipe = ''

cvect_acc = ''

    
### BEGIN SOLUTION
cvect_pipe = Pipeline([('cvect', CountVectorizer()),
                       ('bayes', MultinomialNB())])
cvect_pipe.fit(X_train, y_train)
cvect_acc = cvect_pipe.score(X_test, y_test)
### END SOLUTION


### ANSWER CHECK
cvect_pipe.named_steps
```

Output:
```
{'cvect': CountVectorizer(), 'bayes': MultinomialNB()}
```

<!-- grade_id: cell-ba6a6fa40eeecdc6, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
cvect_pipe_ = Pipeline([('cvect', CountVectorizer()),
                       ('bayes', MultinomialNB())])
cvect_pipe_.fit(X_train, y_train)
cvect_acc_ = cvect_pipe_.score(X_test, y_test)

#
#
#
#
assert cvect_acc == cvect_acc_
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 4

#### Pipeline with `TfidfVectorizer`

**5 Points**

Below, create a pipeline called `tfidf_pipe` with named steps `tfidf` and `bayes` that first vectorizes the text and then uses the `MultinomialNB` estimator with all default settings.  Fit this on the train and score it on the test, assigning the accuracy to `tfidf_acc` below.


<!-- grade_id: cell-7b58933a2378e7c8, points: 0, grade: false, solution: true -->
```python
### GRADED
tfidf_pipe = ''

tfidf_acc = ''

    
### BEGIN SOLUTION
tfidf_pipe = Pipeline([('tfidf', TfidfVectorizer()),
                       ('bayes', MultinomialNB())])
tfidf_pipe.fit(X_train, y_train)
tfidf_acc = tfidf_pipe.score(X_test, y_test)
### END SOLUTION


### ANSWER CHECK
tfidf_pipe.named_steps
```

Output:
```
{'tfidf': TfidfVectorizer(), 'bayes': MultinomialNB()}
```

<!-- grade_id: cell-553b4c71d07d6ac5, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tfidf_pipe_ = Pipeline([('tfidf', TfidfVectorizer()),
                       ('bayes', MultinomialNB())])
tfidf_pipe_.fit(X_train, y_train)
tfidf_acc_ = tfidf_pipe_.score(X_test, y_test)

#
#
#
#
assert tfidf_acc == tfidf_acc_
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 5

#### Assessing performance

**5 Points**

Now, consider searching the hyperparameters of the model.  Specifically, what is the parameter that controls Laplacian smoothing?  Assign your answer as a string to `ans5` below.  As an extra activity, perform a grid search over this parameter and compare the performance to that of `LogisticRegression`.  Also, compare the speed of fit between the logistic and naive bayes models.

<!-- grade_id: cell-95a20d96bbe59b77, points: 0, grade: false, solution: true -->
```python
### GRADED
ans5 = ''

    
### BEGIN SOLUTION
ans5 = 'alpha'
### END SOLUTION

### ANSWER CHECK
print(ans5)
```

Output:
```
alpha
```

<!-- grade_id: cell-17b88d21be4fe45c, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans5_ = 'alpha'

#
#
#
#
assert ans5 == ans5_
### END HIDDEN TESTS
```