### Codio Activity 15.9: Stochastic Gradient Descent

**Expected Time = 30 minutes**

**Total Points = 20**

This activity explores the use of the `SGDRegressor` from scikitlearn.  While there is not a standard gradient descent estimator, the more efficient example of stochastic gradient descent is available.  You will use the earlier credit dataset to explore its use and learn an important lesson about scaling your data with gradient descent methods. 

#### Index

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)


```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.linear_model import SGDRegressor, LinearRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import mean_squared_error
```

```python
credit = pd.read_csv('../../data/Credit.csv', index_col=0)
```

```python
credit.head()
```

```python
X = credit[['Income', 'Limit']]
y = credit['Balance']
```

```python
X_train, X_test, y_train, y_test = train_test_split(X, y, random_state=42)
```

[Back to top](#-Index)

### Problem 1

#### Fitting a basic Linear Regression model

**5 Points**

To begin, use the `LinearRegression` with all default parameters to build a model on the training data and evaluate this on the testing data.  Assign the estimator to `lr` below and training and testing error to `train_mse` and `test_mse` respectively.

<!-- grade_id: cell-070c21d90d3625e3, points: 0, grade: false, solution: true -->
```python
### GRADED
lr = ''
train_mse = ''
test_mse = ''
    
### BEGIN SOLUTION
lr = LinearRegression().fit(X_train, y_train)
train_mse = mean_squared_error(y_train, lr.predict(X_train))
test_mse = mean_squared_error(y_test, lr.predict(X_test))
### END SOLUTION

### ANSWER CHECK
print(train_mse)
print(test_mse)
```

<!-- grade_id: cell-8d7c563dc380a23b, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
lr_ = LinearRegression().fit(X_train, y_train)
train_mse_ = mean_squared_error(y_train, lr_.predict(X_train))
test_mse_ = mean_squared_error(y_test, lr_.predict(X_test))
#
#
#
assert train_mse == train_mse_
assert test_mse == test_mse_
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### Fitting a basic SGD model

**10 Points**

To begin, use the `SGDRegressor` with all default parameters **except** `random_state = 42` to build a model on the training data and evaluate this on the testing data.  Assign the estimator to `sgd_defaults` below.

<!-- grade_id: cell-76a36ca97a2b71c5, points: 0, grade: false, solution: true -->
```python
### GRADED
sgd_defaults = ''
    
    
### BEGIN SOLUTION
sgd_defaults = SGDRegressor(random_state=42).fit(X_train, y_train)
train_mse = mean_squared_error(y_train, sgd_defaults.predict(X_train))
test_mse = mean_squared_error(y_test, sgd_defaults.predict(X_test))
### END SOLUTION

### ANSWER CHECK
print(sgd_defaults)
print(train_mse)
print(test_mse)
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_186/118067456.py in <cell line: 6>()
      4 
      5 ### BEGIN SOLUTION
----> 6 sgd_defaults = SGDRegressor(random_state=42).fit(X_train, y_train)
      7 train_mse = mean_squared_error(y_train, sgd_defaults.predict(X_train))
      8 test_mse = mean_squared_error(y_test, sgd_defaults.predict(X_test))

NameError: name 'X_train' is not defined
```

<!-- grade_id: cell-226f6233b2f17922, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
sgd_defaults_ = SGDRegressor(random_state=42).fit(X_train, y_train)
train_mse_ = mean_squared_error(y_train, sgd_defaults_.predict(X_train))
test_mse_ = mean_squared_error(y_test, sgd_defaults_.predict(X_test))
#
#
#
assert round(train_mse, 0) == round(train_mse_, 0)
assert test_mse == test_mse_
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Scaling the Data

**10 Points**

An important message in gradient descent methods is that scaling the data and using regularization helps to constrain the solution path.  You should be able to see the effect of providing the `SGDRegressor` scaled data below.  Fit a second model using the scaled data below and compare the mean squared error to that of the standard `LinearRegression` estimator.  Assign the estimator with `random_state` to `sgd_scaled` below and the train and test mean squared error to `train_mse` and `test_mse` respectively.

```python
from sklearn.preprocessing import StandardScaler
```

```python
scaler = StandardScaler()
X_tr_scaled = scaler.fit_transform(X_train)
X_ts_scaled = scaler.transform(X_test)
```

<!-- grade_id: cell-a791ed8b60cf0a08, points: 0, grade: false, solution: true -->
```python
### GRADED
sgd_scaled = ''
    
    
### BEGIN SOLUTION
sgd_scaled = SGDRegressor(random_state = 42).fit(X_tr_scaled, y_train)
train_mse = mean_squared_error(y_train, sgd_scaled.predict(X_tr_scaled))
test_mse = mean_squared_error(y_test, sgd_scaled.predict(X_ts_scaled))
### END SOLUTION

### ANSWER CHECK
print(sgd_scaled)
print(train_mse)
print(test_mse)
```

<!-- grade_id: cell-72bfc7f7cfea35b2, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
sgd_scaled_ = SGDRegressor(random_state=42).fit(X_tr_scaled, y_train)
train_mse_ = mean_squared_error(y_train, sgd_scaled_.predict(X_tr_scaled))
test_mse_ = mean_squared_error(y_test, sgd_scaled_.predict(X_ts_scaled))
#
#
#
assert round(train_mse, 0) == round(train_mse_, 0)
assert test_mse == test_mse_
### END HIDDEN TESTS
```

You can return to your earlier examples and see if scaling your data made any differences in the gradient descent convergence, this is important to all models using a gradient descent method and you will see this again with neural networks.