### Full Network Example

This activity focuses on using the full dataset and remaining features with a single layered neural network.  In the last example with only two features an accuracy of roughly 65% was achieved.  Using more features and a similar network architecture you will see if the model improves. 

#### Index 

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
import tensorflow as tf

from sklearn.preprocessing import OneHotEncoder, StandardScaler
from sklearn.compose import make_column_transformer
from sklearn.pipeline import Pipeline
from tensorflow.keras.layers import Dense
from tensorflow.keras.models import Sequential
import warnings
warnings.filterwarnings('ignore')
```

### The Data

Below the titanic data is again loaded.  For this exercise we use columns `['pclass', 'sex', 'age', 'sibsp', 'parch', 'fare', 'embarked', 'class']` as our features to predict `survived`. 

```python
titanic = sns.load_dataset('titanic')
```

```python
X = titanic.loc[:, ['pclass', 'sex', 'age', 'sibsp', 'parch', 'fare', 'embarked', 'class']]
y = titanic['survived']
```

```python
X.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 891 entries, 0 to 890
Data columns (total 8 columns):
 #   Column    Non-Null Count  Dtype   
---  ------    --------------  -----   
 0   pclass    891 non-null    int64   
 1   sex       891 non-null    object  
 2   age       714 non-null    float64 
 3   sibsp     891 non-null    int64   
 4   parch     891 non-null    int64   
 5   fare      891 non-null    float64 
 6   embarked  889 non-null    object  
 7   class     891 non-null    category
dtypes: category(1), float64(2), int64(3), object(2)
memory usage: 49.9+ KB
```

```python
y.head()
```

Output:
```
0    0
1    1
2    1
3    1
4    0
Name: survived, dtype: int64
```

[Back to top](#-Index)

### Problem 1

#### Preparing the features

**10 points**

Below, use the `make_column_transformer` to prepare the features.  Use the `OneHotEncoder` with `drop = 'if_binary'` on all categorical features, and use the `StandardScaler` on the remaining features.  

Be sure to first fill the missing values in `age` with the mean of the column.

Assign the transformed array to `X_t` below.

<!-- grade_id: cell-537339f5cc20a045, points: 0, grade: false, solution: true -->
```python
### GRADED
X_t = ''


    
### BEGIN SOLUTION
X['age'] = X['age'].fillna(X['age'].mean())


t = make_column_transformer((OneHotEncoder(drop='if_binary'),['sex', 'embarked', 'class']),
                           remainder=StandardScaler())
X_t = t.fit_transform(X)
### END SOLUTION

### ANSWER CHECK
X_t.shape
```

Output:
```
(891, 13)
```

<!-- grade_id: cell-9c455d637c5f8658, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
X_ = titanic.loc[:, ['pclass', 'sex', 'age', 'sibsp', 'parch', 'fare', 'embarked', 'class']]
y_ = titanic['survived']
X_['age'] = X_['age'].fillna(X_['age'].mean())


t_ = make_column_transformer((OneHotEncoder(drop='if_binary'),['sex', 'embarked', 'class']),
                           remainder=StandardScaler())
X_t_ = t_.fit_transform(X_)

#
#
#
np.testing.assert_array_equal(X_t, X_t_)
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### The Network Architecture

**10 points**

Below, construct a network named `model` that has one hidden layer with 100 Dense units that use the `relu` activation function.  Use an output layer with one node that uses the `sigmoid` activation function.  

<!-- grade_id: cell-c56c772a18ea7232, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
model = ''

   
### BEGIN SOLUTION
tf.random.set_seed(42)
model = Sequential([
    Dense(100, activation = 'relu'),
    Dense(1, activation = 'sigmoid')
])
### END SOLUTION

### ANSWER CHECK
model
```

Output:
```
<keras.engine.sequential.Sequential at 0x7fd60e937160>
```

<!-- grade_id: cell-9c5bd1cdd35a6125, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
model_ = Sequential([
    Dense(100, activation = 'relu'),
    Dense(1, activation = 'sigmoid')
])

#
#
#
assert model_.layers[0].activation == model.layers[0].activation
assert model_.layers[1].activation == model.layers[1].activation
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Train and Evaluate the Network

**10 points**

Finally, train and evaluate the network using the following compilation settings.  Assign the fit model to `history` below.

- `optimizer = 'rmsprop'`
- `loss = 'bce'`
- `metrics = ['accuracy']`
- `epochs = 20`
- `batch_size = 10`
- `verbose = 0`

Also, be sure to leave the `tf.random.set_seed(42)` for proper grading.

<!-- grade_id: cell-d0ef12fa2d67be00, points: 0, grade: false, solution: true -->
```python
### GRADED
history = ''

   
### BEGIN SOLUTION
tf.random.set_seed(42)
model = Sequential([
    Dense(100, activation = 'relu'),
    Dense(1, activation = 'sigmoid')
])
model.compile(loss = 'bce', optimizer = 'rmsprop',
                     metrics = ['accuracy'])
np.random.seed(42)
history = model.fit(X_t, y, epochs = 20, batch_size = 10,
                           verbose = 0)
### END SOLUTION

### ANSWER CHECK
history.history['accuracy'][-1]
```

Output:
```
0.8428731560707092
```

<!-- grade_id: cell-39e494f7d324b681, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
model_ = Sequential([
    Dense(100, activation = 'relu'),
    Dense(1, activation = 'sigmoid')
])
model_.compile(loss = 'bce', optimizer = 'rmsprop',
                     metrics = ['accuracy'])
np.random.seed(42)
history_ = model_.fit(X_t_, y, epochs = 20, batch_size = 10,
                           verbose = 0)

#
#
#
assert history_.history['accuracy'][-1] == history.history['accuracy'][-1]
### END HIDDEN TESTS
```