### Codio Activity 6.3: Plotting and Interpreting Singular Values

**Expected Time: 60 Minutes**

**Total Points: 20 Points**

This activity focuses on building a scree plot with Python.  A scree plot can be used to determine how many Principal Components to keep in a given dataset.  Using an `sklearn` built in housing dataset, you will plot the values of $\Sigma$ in decreasing order -- this is a scree plot.  For more on scree plots see [here](https://en.wikipedia.org/wiki/Scree_plot#:~:text=In%20multivariate%20statistics%2C%20a%20scree,principal%20component%20analysis%20(PCA).).

## Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
from sklearn.datasets import fetch_california_housing
import seaborn as sns
from scipy.linalg import svd
from sklearn.datasets import make_regression
```

```python
housing = fetch_california_housing(as_frame = True, data_home='data')
```

```python
df = housing.frame.head(50).drop('MedHouseVal', axis = 1) #limits due to computation 
```

```python
df.head()
```

Output:
```
MedInc  HouseAge  AveRooms  AveBedrms  Population  AveOccup  Latitude  \
0  8.3252      41.0  6.984127   1.023810       322.0  2.555556     37.88   
1  8.3014      21.0  6.238137   0.971880      2401.0  2.109842     37.86   
2  7.2574      52.0  8.288136   1.073446       496.0  2.802260     37.85   
3  5.6431      52.0  5.817352   1.073059       558.0  2.547945     37.85   
4  3.8462      52.0  6.281853   1.081081       565.0  2.181467     37.85   

   Longitude  
0    -122.23  
1    -122.22  
2    -122.24  
3    -122.25  
4    -122.25
```

```python
df.shape
```

Output:
```
(50, 8)
```

[Back to top](#Index:) 

## Problem 1

### Extracting $\Sigma$

**10 Points**

Complete the function `singular_values` according to the instructions below:

- The function takes, as input, an aray `X`, and a booloean flag `scale`, describing whether you want to scale the array `X` or not.
- Using an `if` statement, check if `scale` is `True`. If it is, scale `X` using its mean and standard deviation.
- If `scale` is not `True`, use the `svd` function to factor `X` into  `u`, `sigma`, and `vt`.
- The function should return the numpy array of singular values of X, `sigma`.


<!-- grade_id: cell-4190c44417f23648, points: 0, grade: false, solution: true -->
```python
### GRADED

def singular_values(X, scale = False):
    """Return the singular values resulting from 
    SVD decomposition.  

    Parameters
    ----------
    X: np.array or pd.DataFrame
        An array of data
    scale: boolean
        Boolean determines whether data needs to be scaled

    Returns an numpy array of singular values of X
    """
    # YOUR CODE HERE 
    return None # replace with your return statement

### BEGIN SOLUTION
def singular_values(X, scale = False):
    if scale:
        X = (X - X.mean())/X.std()
        
    u, sigma, vt = svd(X)
    return sigma
### END SOLUTION

# Answer check
print(type(singular_values(df)))
sigma = singular_values(df)
print(sigma.shape)
```

Output:
```
<class 'numpy.ndarray'>
(8,)
```

<!-- grade_id: cell-5474a2fc14639bda, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def singular_values_(X, scale = False):
    if scale:
        X = (X - X.mean())/X.std()
        
    u, sigma, vt = svd(X)
    return sigma

sigma1 = singular_values(df)
sigma2 = singular_values(df, scale = True)
sigma1_ = singular_values_(df)
sigma2_ = singular_values_(df, scale = True)
#
#
#
np.testing.assert_array_equal(sigma1, sigma1_, 'Make sure not scaling works')
np.testing.assert_array_equal(sigma2, sigma2_, 'Make sure your scaling works')
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 2

### Plotting $\Sigma$

**10 Points**

Using the function above, build a plot of the singular values. Use your plot to determine how many principal components you should keep from the dataframe `df`.  Assign your solution as an integer to `ans2` below.



<!-- grade_id: cell-81c203e4b436cbe6, points: 0, grade: false, solution: true -->
```python
### GRADED

#plot of singular values here

#how many components should you keep?
ans2 = ''

### BEGIN SOLUTION
ans2 = 2
### END SOLUTION
```

<!-- grade_id: cell-5029bf3266576b52, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans2_ = [1, 2, 3, 4, 5]
#
#
#
assert ans2 in ans2_, 'Make sure your solution is an integer.'
### END HIDDEN TESTS
```

Depending on the dataset, we may have a different number of principal components to keep.  While not a hard fast statistical test, using a scree plot can provide some information on the appropriate number of values to keep.  

```python

```