### Codio Activity 23.4: Fine Tuning a Pre-trained Network

**Expected Time = 60 minutes**

**Total Points = 30**

In addition to the use of a pre-trained network to extract features from a different dataset, the weights can be adjusted or **fine-tuned** as a last step to squeeze additional performance from the network.  To do so, you will again use the `EfficientNetV2B0` network on the `cifar10` data from `keras`.  This time you are encouraged to use the functional API syntax to construct your network.  

For a second example, consult the `keras` documentation example [here](https://keras.io/guides/transfer_learning/). 

#### Index

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)

Run the code cell below to import the necessary libraries.

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt


from tensorflow.keras.applications.efficientnet_v2 import EfficientNetV2B0
from tensorflow.keras.layers import Dense, Flatten
from tensorflow.keras.models import Sequential
from tensorflow.keras.utils import to_categorical
import tensorflow as tf
from tensorflow.keras.datasets import cifar10
from tensorflow.keras import Model, Input
```

Output:
```
2022-11-21 16:58:50.322351: W tensorflow/stream_executor/platform/default/dso_loader.cc:64] Could not load dynamic library 'libcudart.so.11.0'; dlerror: libcudart.so.11.0: cannot open shared object file: No such file or directory
2022-11-21 16:58:50.322409: I tensorflow/stream_executor/cuda/cudart_stub.cc:29] Ignore above cudart dlerror if you do not have a GPU set up on your machine.
```

```python
(X_train, y_train), (X_test, y_test) = cifar10.load_data()
sub_indices = pd.DataFrame(y_train, columns= ["labels"]).groupby('labels', group_keys=False).apply(lambda x: x.sample(frac=0.3, random_state=123)).index
X_train = X_train[sub_indices]
y_train = y_train[sub_indices]

Y_train = to_categorical(y_train)
Y_test = to_categorical(y_test)

```

[Back to top](#-Index)

### Problem 1

#### Training the Network to Convergence

**10 Points**

In the code cell below, use the function `EfficientNetV2B0` with the appropriate input shape and the argument `include_top` equal to `False` to create the model you will be using for this activity. Assign your result to `base_model`.

Use the function `Input()` with argument `shape` equal to `(32, 32, 3)` and assign your result to the variable `inputs`.

Use `base_model` with argument equal to `inputs` and assign your result to the variable `x`.

Use the function `Flatten` to flatten `x`. The pseudocode to complete this step is given below:

```Python
x = Flatten()(...)
```

Pass `x` through a `Dense` layer with 10 hidden nodes and `activation` equal to `relu` and assign the result to the variable `output`. The pseudocode to complete this step is given below:

```Python
output = Dense(..., activation = ...)(...)
```

Use the code `base_model.trainable = False` to ensure that your weights are not trainable.

Use the function `Model()` with argument `inputs` and `output` to define your model. Assign the result to the variable `model`.

Compile `model` using `categorical_crossentropy` as your `loss` and `accuracy` as your `metric`.

Use the `fit()` function on `model` to fit the  training data `X_train` and `Y_train`. Set the argument `validation_data` equal to `(X_test, Y_test)` and the argument `epochs` equal to 1.  Assign the result to the variable `bottom_model` below. 

NOTE: This question is computationally expensive, so please be patient with the processing. It may take a few minutes based on your computing power. 


<!-- grade_id: cell-e92e20d0c85d3026, points: 0, grade: false, solution: true -->
```python
### GRADED
base_model = ''
inputs = ''
x = ''
x = ''
x = ''
output = ''
model = ''
#be sure to compile

bottom_model = ''
    
### BEGIN SOLUTION
base_model = EfficientNetV2B0(input_shape = (32, 32, 3), include_top=False)
inputs = Input(shape = (32, 32, 3))
x = base_model(inputs)
x = Flatten()(x)
x = Dense(10, activation = 'relu')(x)
output = Dense(10, activation = 'sigmoid')(x)
base_model.trainable = False
model = Model(inputs, output)
model.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
bottom_model = model.fit(X_train, Y_train, validation_data=(X_test, Y_test), 
                    epochs = 1)
### END SOLUTION

### ANSWER CHECK
print(model.summary())
```

Output:
```
2022-11-21 16:59:03.831934: W tensorflow/stream_executor/platform/default/dso_loader.cc:64] Could not load dynamic library 'libcuda.so.1'; dlerror: libcuda.so.1: cannot open shared object file: No such file or directory
2022-11-21 16:59:03.831995: W tensorflow/stream_executor/cuda/cuda_driver.cc:269] failed call to cuInit: UNKNOWN ERROR (303)
2022-11-21 16:59:03.833111: I tensorflow/stream_executor/cuda/cuda_diagnostics.cc:156] kernel driver does not appear to be running on this host (stronggarbo-aztecimmune): /proc/driver/nvidia/version does not exist
2022-11-21 16:59:03.834818: I tensorflow/core/platform/cpu_feature_guard.cc:193] This TensorFlow binary is optimized with oneAPI Deep Neural Network Library (oneDNN) to use the following CPU instructions in performance-critical operations:  AVX2 FMA
To enable them in other operations, rebuild TensorFlow with the appropriate compiler flags.
Epoch 1/5
469/469 [==============================] - 50s 88ms/step - loss: 1.5093 - accuracy: 0.467
```

<!-- grade_id: cell-c7a974852163092b, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
base_model_ = EfficientNetV2B0(input_shape = (32, 32, 3), include_top=False)
inputs_ = Input(shape = (32, 32, 3))
x_ = base_model(inputs_)
x_ = Flatten()(x_)
x_ = Dense(100, activation = 'relu')(x_)
output_ = Dense(10, activation = 'sigmoid')(x_)
model_ = Model(inputs_, output_)
model_.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
#
#
#
assert type(base_model) == type(base_model_)
### END HIDDEN TESTS
print(model_.summary())
```

Output:
```
Epoch 1/5
469/469 [==============================] - 52s 90ms/step - loss: 1.5343 - accuracy: 0.4644 - val_loss: 1.2853 - val_accuracy: 0.5555
Epoch 2/5
469/469 [==============================] - 30s 65ms/step - loss: 1.2563 - accuracy: 0.5579 - val_loss: 1.2166 - val_accuracy: 0.5710
Epoch 3/5
469/469 [==============================] - 32s 68ms/step - loss: 1.1789 - accuracy: 0.5853 - val_loss: 1.1935 - val_accuracy: 0.5897
Epoch 4/5
469/469 [==============================] - 35s 74ms/step - loss: 1.1236 - accuracy: 0.6075 - val_loss: 1.1563 - val_accuracy: 0.5891
Epoch 5/5
469/469 [==============================] - 34s 73ms/step - loss: 1.0934 - accuracy: 0.6144 - val_loss: 1.2450 - val_accuracy: 0.5783
Model: "model_1"
_________________________________________________________________
 Layer (type)                Output Shape              Param #   
=================================================================
 input_4 (InputLayer)        [(None, 32, 32, 3)]       0         
    
```

[Back to top](#-Index)

### Problem 2

#### Setting to Trainable

**10 Points**

In the code cell below, use the code `base_model.trainable = True` to ensure that your weights are now trainable.


Set the final five layers to trainable in the `base_model` as demonstrated in the lectures. 

```python
base_model.trainable #are the base model weights set to trainable?
```

Output:
```
False
```

<!-- grade_id: cell-17e4f0a4de6dc3b8, points: 0, grade: false, solution: true -->
```python
### GRADED
base_model.trainable = True
for layer in '':
    #make trainable
    pass
    
### BEGIN SOLUTION
base_model.trainable = True
for layer in base_model.layers[:-5]:
    layer.trainable = False
### END SOLUTION

### ANSWER CHECK
for i, layer in enumerate(base_model.layers[-10:]):
    print(f'Layer is trainable: {layer.trainable}')
```

Output:
```
Layer is trainable: False
Layer is trainable: False
Layer is trainable: False
Layer is trainable: False
Layer is trainable: False
Layer is trainable: True
Layer is trainable: True
Layer is trainable: True
Layer is trainable: True
Layer is trainable: True
```

<!-- grade_id: cell-279935493600bece, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
base_model_.trainable = True
for layer in base_model_.layers[:-5]:
    layer.trainable = False
#
#
#
for i, layer in enumerate(base_model.layers[-5:]):
    assert layer.trainable == True
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Refitting the network

**10 Points**

In the code cell below, use the function `compile` on `model` using `categorical_crossentropy` as your `loss` and `accuracy` as your `metric`.

Next, use the `fit()` function on `model` to fit the  training data `X_train` and `Y_train`. Set the argument `validation_data` equal to `(X_test, Y_test)` and the argument `epochs` equal to 1.  Assign the result to the variable `fine_tuned_history` below. 
 

<!-- grade_id: cell-75f43fbafe9c0840, points: 0, grade: false, solution: true -->
```python
### GRADED
fine_tuned_history = ''
    
### BEGIN SOLUTION
model.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
fine_tuned_history = model.fit(X_train, Y_train, validation_data = (X_test, Y_test), epochs = 1)
### END SOLUTION

### ANSWER CHECK
fine_tuned_history.history['accuracy'][-1]
```

Output:
```
Epoch 1/5
469/469 [==============================] - 56s 90ms/step - loss: 1.6604 - accuracy: 0.5090 - val_loss: 1.2255 - val_accuracy: 0.5940
Epoch 2/5
469/469 [==============================] - 40s 85ms/step - loss: 1.0498 - accuracy: 0.6427 - val_loss: 1.1942 - val_accuracy: 0.6157
Epoch 3/5
469/469 [==============================] - 39s 82ms/step - loss: 0.9001 - accuracy: 0.6948 - val_loss: 1.1988 - val_accuracy: 0.6255
Epoch 4/5
469/469 [==============================] - 44s 95ms/step - loss: 0.7925 - accuracy: 0.7295 - val_loss: 1.3052 - val_accuracy: 0.6044
Epoch 5/5
469/469 [==============================] - 36s 77ms/step - loss: 0.7065 - accuracy: 0.7631 - val_loss: 1.3319 - val_accuracy: 0.6179
0.7630666494369507
```

<!-- grade_id: cell-c60ddea5358927c2, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
model_.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
#
#
#
#assert len(fine_tuned_history.history['accuracy']) == len(fine_tuned_history_.history['accuracy'])
### END HIDDEN TESTS
```

Output:
```
Epoch 1/5
469/469 [==============================] - 55s 94ms/step - loss: 0.9569 - accuracy: 0.6874 - val_loss: 1.3558 - val_accuracy: 0.6133
Epoch 2/5
469/469 [==============================] - 35s 76ms/step - loss: 0.7281 - accuracy: 0.7573 - val_loss: 1.3930 - val_accuracy: 0.6100
Epoch 3/5
469/469 [==============================] - 35s 74ms/step - loss: 0.6363 - accuracy: 0.7869 - val_loss: 1.4276 - val_accuracy: 0.6243
Epoch 4/5
469/469 [==============================] - 36s 77ms/step - loss: 0.5828 - accuracy: 0.8080 - val_loss: 1.5683 - val_accuracy: 0.6173
Epoch 5/5
469/469 [==============================] - 34s 73ms/step - loss: 0.5433 - accuracy: 0.8201 - val_loss: 1.5925 - val_accuracy: 0.6181
```